**T&B Exercise 8.2.** Write a julia function `Q,R = mgs(A)` that computes a reduced QR factorization $A = \hat Q \hat R$ of an $m \times n$ matrix $A$ with $m \geq n$ using modified Gram-Schmidt orthogonalization.  The output variables are a matrix $Q \in \mathbb{C}^{m \times n}$ orthonormal columns and a triangular matrix $R \in \mathbb{C}^{n \times n}$.

```julia
using LinearAlgebra
```

<!-- grade_id: cell-c473ca965ce3c8a9, points: 0, grade: false, solution: true -->
```julia
function mgs(A)
    n = size(A)[2]
    Q = copy(A)
    R = zeros(eltype(A), n, n)
    for i in 1:n
        # precondition Q[:, i] is orthogonal
        R[i, i] = norm(Q[:, i])
        Q[:, i] /= R[i, i]
        R[i, i+1:end] = Q[:, i]' * Q[:, i+1:end]
        Q[:, i+1:end] -= Q[:, [i]] * R[[i], i+1:end]
    end
    return Q, R
end
```

Output:
```
mgs (generic function with 1 method)
```

<!-- grade_id: cell-d366f91037f0f92f, points: 1, grade: true, solution: false -->
```julia
m = 100
n = 63

A = randn(m,n)
Q,R = mgs(A)
orthog_err = norm(I - Q'*Q,2)
recon_err = norm(A - Q*R,2)
triangle_err = norm(R - UpperTriangular(R),2)
println("Modified Gram-Schmidt:")
println("  orthogonality error: $orthog_err")
println("  reconstruction error: $recon_err")
println("  triangularity error: $triangle_err")
@assert triangle_err == 0
@assert orthog_err < 1.e-12
@assert recon_err < 1.e-12
```

Output:
```
Modified Gram-Schmidt:
  orthogonality error: 3.809795701115102e-15
  reconstruction error: 2.080447942559563e-14
  triangularity error: 0.0
```